## Assignment 1.2 Python Classes
The constructor is always written as a function called init() It must always take as its first argument a reference to the instance being constructed Single underscore before a member variable denotes it protected and double underscore denotes it private but it is just a convention.

### Excersise 3
Modify rectangle class to compare rectangle objects (comparing with areas) 
Operator overloading of rich comparisons = ,> ,< ,<= ,>= ,!= , 


<!-- grade_id: cell-57c54f6b2e9f1bdd, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Rectangle():
    def __init__(self, length, breadth):
        self.length = length
        self.breadth = breadth
    def get_area(self):
        return self.length*self.breadth
    def get_perimeter(self):
        return 2*(self.length+self.breadth)
    def __repr__(self):
        return "Rectangle(length={}, breadth={})".format(self.length,self.breadth)
    def __lt__(self, other):
        return self.get_area() < other.get_area()
    def __le__(self, other):
        return self.get_area() <= other.get_area()
    def __eq__(self, other):
        return self.get_area() == other.get_area()
    def __ge__(self, other):
        return self.get_area() >= other.get_area()
    def __gt__(self, other):
        return self.get_area() > other.get_area()
    def __ne__(self, other):
        return self.get_area() != other.get_area()    
### END SOLUTION
```

<!-- grade_id: cell-2a7ef91ac3825c4f, points: 5, grade: true, solution: false -->
```python
assert(Rectangle(3,10)==Rectangle(3,10))
assert(Rectangle(2,10)<Rectangle(3,10))
assert(Rectangle(3,10)!=Rectangle(2,10))
assert(Rectangle(3,10)>=Rectangle(3,10))
### HIDDEN TESTS HERE ###
### BEGIN HIDDEN TESTS
assert(Rectangle(3,10)<=Rectangle(3,10))
assert(Rectangle(3,10)>Rectangle(2,10))
### END HIDDEN TESTS
```

### Excersise 4
Inheritance 
Derive a class **Square** from **Rectangle** with one properties "length" (initialized through constructor) and able to calculate area, perimeter and comparison

<!-- grade_id: cell-9193fb70b8589235, points: 0, grade: false, solution: true -->
```python
class Square():
    pass
### BEGIN SOLUTION
class Square(Rectangle):
    def __init__(self, length):
            self.length = length
            self.breadth = self.length
### END SOLUTION
```

<!-- grade_id: cell-b860a4edf7a89cec, points: 5, grade: true, solution: false -->
```python
assert(Square(9).get_area()==81)
assert(Square(10) != Square(5))
### HIDDEN TESTS HERE ###
### BEGIN HIDDEN TESTS
assert(Square(10).get_area()==100)
assert(Square(10)>Square(5))
assert(isinstance(Square,type(Rectangle)))
### END HIDDEN TESTS
```