## Assignment 1.2 Python
Python is an easy to learn, powerful programming language with efficient high-level data structures and object-oriented programming. Python’s elegant syntax and dynamic typing, together with its interpreted nature, make it an ideal language for scripting and rapid application development in many areas on most platforms.

### Exercise 3
Write a program which will list all such numbers which are divisible by 7 but are not a multiple of 5, between 2000 and 3200 (both included).logical Operator, for, if, and range could be helpful.

<!-- grade_id: cell-ab0fa55a5b37052d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84dacac197b14f03, points: 5, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ####

```

### Exercise 4 (Function)
write a function fibbo to return a list of "n" (greater than 2) fibonnaci number

<!-- grade_id: cell-525cbf27d5e3196d, points: 0, grade: false, solution: true -->
```python
def fibbo(n):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-964b4920f43ba78f, points: 5, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ###
assert(fibbo(5)==[1, 1, 2, 3, 5])
```