## Assignment 1.2 Python Classes
The constructor is always written as a function called init()
It must always take as its first argument a reference to the instance being constructed
Single underscore before a member variable denotes it protected and double underscore denotes it private but it is just a convention.

### Excercise1 
Create a Class Rectangle such that its objects are initialized with its length and breadth. Make two class methods get_area and get_perimeter which calculates and return the area and circumference respectively


<!-- grade_id: cell-88007790521d6ed4, points: 0, grade: false, solution: true -->
```python
class Rectangle():
# Your code here
    def get_area():
        # Code here
        pass
    def get_perimeter():
        #Code here
        pass
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-01c4e2988b52334e, points: 5, grade: true, solution: false -->
```python
rec1 = Rectangle(4,10)
assert(rec1.get_area() == 40)
### HIDDEN TESTS HERE ###
```

### Excersie2 
Python Special Methods

<!-- grade_id: cell-648325934ae54185, points: 0, grade: false, solution: true -->
```python
# Modify the Rectangle class to print object as "Rectangle(length=a, breadth=b)"
# where a and b are length and breadth 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eab820ca1546be91, points: 5, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ###
```