## Assignment 1.2 Python Classes
The constructor is always written as a function called init()
It must always take as its first argument a reference to the instance being constructed
Single underscore before a member variable denotes it protected and double underscore denotes it private but it is just a convention.

### Excercise1 
Create a Class Rectangle such that its objects are initialized with its length and breadth. Make two class methods get_area and get_perimeter which calculates and return the area and circumference respectively


<!-- grade_id: cell-88007790521d6ed4, points: 0, grade: false, solution: true -->
```python
class Rectangle():
# Your code here
    def get_area():
        # Code here
        pass
    def get_perimeter():
        #Code here
        pass
### BEGIN SOLUTION
class Rectangle():
    def __init__(self, length, breadth):
        self.length = length
        self.breadth = breadth
    def get_area(self):
        return self.length*self.breadth
    def get_perimeter(self):
        return 2*(self.length+self.breadth)
### END SOLUTION
```

<!-- grade_id: cell-01c4e2988b52334e, points: 5, grade: true, solution: false -->
```python
rec1 = Rectangle(4,10)
assert(rec1.get_area() == 40)
### HIDDEN TESTS HERE ###
### BEGIN HIDDEN TESTS
rec91 = Rectangle(3,10)
assert(rec91.get_area() == 3*10)
assert(rec91.get_perimeter() == 2*(3+10))
### END HIDDEN TESTS
```

### Excersie2 
Python Special Methods

<!-- grade_id: cell-648325934ae54185, points: 0, grade: false, solution: true -->
```python
# Modify the Rectangle class to print object as "Rectangle(length=a, breadth=b)"
# where a and b are length and breadth 

### BEGIN SOLUTION
class Rectangle():
    def __init__(self, length, breadth):
        self.length = length
        self.breadth = breadth
    def get_area(self):
        return self.length*self.breadth
    def get_perimeter(self):
        return 2*(self.length+self.breadth)
    def __repr__(self):
        return "Rectangle(length={}, breadth={})".format(self.length,self.breadth)
### END SOLUTION
```

<!-- grade_id: cell-eab820ca1546be91, points: 5, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ###
### BEGIN HIDDEN TESTS
rec91 = Rectangle(3,10)
assert(rec91.__repr__()=="Rectangle(length=3, breadth=10)")
### END HIDDEN TESTS
```