## Assignment 1.1 Python
Python is an easy to learn, powerful programming language with efficient high-level data structures and object-oriented programming. Python’s elegant syntax and dynamic typing, together with its interpreted nature, make it an ideal language for scripting and rapid application development in many areas on most platforms.


### Exercise 1
T is a Nested Tuple, use indexing to extract 'g' from the Tuple

<!-- grade_id: cell-bc0474cecedb5d8b, points: 0, grade: false, solution: true -->
```python
T = ('a','b','c',('d','e'),('f',('f','g'),'k'),'p')
thisisg = None
### BEGIN SOLUTION
thisisg = T[4][1][1]
### END SOLUTION
```

<!-- grade_id: cell-d4a09f8a55d85bbc, points: 5, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ####
assert(thisisg=='g')
### BEGIN HIDDEN TESTS
assert(thisisg=='g')
### END HIDDEN TESTS
```

### Exercise 2 (Lists)
Find the sum of the smallest 10 numbers of List L Indexing, Sorted and sum function could be useful

<!-- grade_id: cell-1b1563dfa6a1140e, points: 0, grade: false, solution: true -->
```python
L =[75, 28, 47, 51, 96, 76, 32, 57, 10, 10, 25, 62, 17,  9, 68, 65, -2,
       54, 34, 41, 74, 83, 91, 59, -3, 88, 59, 91, 10, -9, 90, 46, 58, 25,
       62, 32, 74,  1, 61, 43, 25, 62, -5, 49, 40, 44, 83, -5, 13, 35]

### BEGIN SOLUTION
ans = sum(sorted(L)[:10])
### END SOLUTION
```

<!-- grade_id: cell-b939687d9fb4af3a, points: 5, grade: true, solution: false -->
```python
### HIDDEN TESTS HERE ####
assert(ans==16)
### BEGIN HIDDEN TESTS
assert(ans==sum(sorted(L)[:10]) )
### END HIDDEN TESTS
```